<problem display_name="Quantum mechanics ket input" attempts="1000" weight="10">
  <text>
    <p>
Quantum mechanics uses &#x201C;ket" notation like [mathjaxinline]|\psi \rangle[/mathjaxinline] to indicate unit vectors. </p>
    <p>
This problem demonstrates how to use customresponse python code to check for equality of formulas with such ket notation, as well as a javascript pre-processor to properly format ket expressions. </p>
    <p>
Type in <tt class="tt">(2*|0&gt;+3*|b&gt;)/sqrt(13)</tt> as the answer below. </p>
    <p>
Multiple correct answer are supported. <tt class="tt">|b&gt;</tt> is also correct, in this example. </p>
    <font color="red">
      <b> Remeber to explicitly denote multiplication using [mathjaxinline]\tt *[/mathjaxinline]</b>
    </font>
    <p>
      <br/>
    </p>
    <script type="text/python" system_path="python_lib">

from ket_evaluator import *

</script>
    <p>
      <p style="display:inline">[mathjaxinline]|\lambda \rangle  =[/mathjaxinline] </p>
      <customresponse cfn="test_ket" inline="1" expect="(2*|0&amp;gt;+3*|b&amp;gt;)/sqrt(13)" options="samples='a,b@1,1:20,20#50'!altanswer='|a&amp;gt;'" cfn_extra_args="options">
        <textline size="60" correct_answer="(2*|0&amp;gt;+3*|b&amp;gt;)/sqrt(13)" inline="1" math="1" preprocessorClassName="MathjaxPreprocessorForQM" preprocessorSrc="/static/js/mathjax_preprocessor_for_QM.js"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
The [mathjaxinline]\hat{n}[/mathjaxinline] vector goes through the point [mathjaxinline](\theta , \phi )[/mathjaxinline] specified in polar coordinates, with [mathjaxinline]\phi =0[/mathjaxinline] and </p>
          <table id="a0000000006" cellpadding="7" width="100%" cellspacing="0" class="eqnarray" style="table-layout:auto">
            <tr id="a0000000007">
              <td style="border:none">&#xA0;</td>
              <td style="border:none">
	    [mathjaxinline]\displaystyle \theta = \cos ^{-1} \frac{\omega _0/2}{\sqrt {V^2 + \left(\frac{\omega _0}{2}\right)^2}} = \tan ^{-1}\left(\frac{2V}{\omega _0}\right) \, .[/mathjaxinline]
        </td>
              <td style="border:none">&#xA0;</td>
              <td style="border:none" class="eqnnum">
                <span>(<span>4</span>)</span>
              </td>
            </tr>
          </table>
          <p>
In terms [mathjaxinline]\theta[/mathjaxinline], the eigenstates may be expressed as </p>
          <table id="a0000000008" cellpadding="7" width="100%" cellspacing="0" class="eqnarray" style="table-layout:auto">
            <tr id="a0000000009">
              <td style="border:none">&#xA0;</td>
              <td style="border:none">
	    [mathjaxinline]\displaystyle |\lambda _+\rangle  = \cos \frac{\theta }{2} |a\rangle  +\sin \frac{\theta }{2} |b\rangle[/mathjaxinline]
        </td>
              <td style="border:none">&#xA0;</td>
              <td style="border:none" class="eqnnum">
                <span>(<span>5</span>)</span>
              </td>
            </tr>
            <tr id="a0000000010">
              <td style="border:none">&#xA0;</td>
              <td style="border:none">
	    [mathjaxinline]\displaystyle |\lambda _-\rangle  = -\sin \frac{\theta }{2} |a\rangle  +\cos \frac{\theta }{2} |b\rangle  \, .[/mathjaxinline]
        </td>
              <td style="border:none">&#xA0;</td>
              <td style="border:none" class="eqnnum">
                <span>(<span>6</span>)</span>
              </td>
            </tr>
          </table>
        </font>
      </solution>
    </p>
  </text>
</problem>
